# Introduction to Python for engineers


Around 7 years ago I have come across this beautiful language and fell in love with it! It helped me with my classes, my PhD career, my teaching, and many other things. 

Just as every spoken language consists of words and sentences, a programming language consists of variables and operators. We first will learn what they are, and secondly we will learn why we need them. Then we will start constructing our first "sentences" - in other words - writing programs to solve our engineering problems.

Here are the main elements that we will learn today.

1. Python syntaxes
2. Arithmetic operations in Python
3. Conditionals
4. Loops
5. Functions

These are the basic elements of every language. Don't consider them to be your goal, these are the tools. The best thing you can do with tools is to use them. So go ahead and start solving the problems right away! 


## Variables and arithmetic operations



```python
# the first thing that you will need in Python are variables
# You can assign a value to them and perform arithmetic operations with them
a = 1
b = 2
```

```python
# If you want to see their values you can use the `print()` function

print(a)
```

```python
#here is how to perform simple arithmetic operations with those variables

c = a - b 
print(c)
```

```python
# going further
d = a*b
e = a/b
f = 9**(0.5)
# if you want to print several variables, just list them with a comma
print(d, e, f)
```

```python
#NOTE:

#python variables can have different type:
# integer:
a1 = 1
# float
f1 = 1.0
#string
s1 = "1"

#although for us it is the same number, these are different variables for Python
print(type(a1))
print(type(f1))
print(type(s1))

```

### Project 1

Enough discussion lets do our first little project. Lets calculate the area of a triangle with sides (x = 3, y = 4, z = 5) using the formula:

1. Using Phales's theorem $A = \sqrt{s\cdot(s-x)(x-y)(s-z)}$, where $s = x + y + z$


```python
# First lets assign values to variables 

x = 3
y = 4
z = 5

#lets calculate the perimeter of the triangle 
s = (x+y+z)/2.0

# now using s, lets calculate the area
A = (s*(s-x)*(s-y)*(s-z))**0.5
```

```python
print(A)
```

### Project 2

Calculate the area of a triangle if its base $b = 30$, its height $h = 40$

<!-- grade_id: cell-d86341091dae4c1a, points: 0, grade: false, solution: true -->
```python
def area():
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-c8dec45d13b114cc, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(area(),600)
```

### Project 3

Calculate the area of a trapezoid if its base $b1 = 5$, its height $h = 2$ and the length of the top side $a = 3$.

<!-- grade_id: cell-61200295d52138a8, points: 0, grade: false, solution: true -->
```python
def area_trapz():

    #area of a trapz
    a1 = 3
    b1  = 5
    h = 2
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1a7bd9ce0076c6f5, points: 1, grade: true, solution: false -->
```python
assert_equal(area_trapz(),8)
```

### Project 4

Lets get a problem from the world closer to engineering.
Calculate the pressure of two moles of ideal gas, at temperature $T = 27 C$, and volume $V = 1 litre$

Take zero Kelvin as being $-273$ degrees Celcius and NOT $-273.15$ degrees Celcius

<!-- grade_id: cell-ecd7a114d1a4c62c, points: 0, grade: false, solution: true -->
```python
#CHBE problem
V = 1 #l
R = 8.31 #J/(mol K)
T = 27 #C
n = 1

def two_moles():

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fdfeea03e310d99f, points: 1, grade: true, solution: false -->
```python
assert_equal(two_moles(),2493000)
```

## Conditionals

Sometimes we want to compare values and proceed if only a certain condition is met. In these cases we use `if ... elseif ... else` statements.
But before we start, let me introduce you to another data type - Boolean.
It can have only two values `True, False`, but it is crucial for the design of many important algorithms in the modern computer science.


```python
# the boolean operators can be a result of conditional operations 
# 2 >= 1 = True
# For example
if 1>2:
    print("This is false")
```

```python
if 2<1:
    print("I am in the if")
elif 0>2:
    print("I am in the elif")
else:
    print("I am in else")

```

```python
# Boolean operators can be "multiplied/added" with and/or
# True and False = False
# True and True = True
# False and False = False
# True or False = True
# True or True = True
# False or False = False

if (True) or (False):
    print("I will no go here")
if (True) and (False):
    print("I will no go here")
    
```

One reason why I love Python is its flexibility. Unlike Matlab that comes with lots of packages pre-installed and important, he core of of Python doesn't have virtually anything.
To get the full power of Python we have to import libraries that we will need.
To do that we use the `import` operator.

`import random` imports the library for generating random numbers.


```python
#random things
import random
r = random.randint(1,10)
print(r)
```

### Project 5

Imagine you play a game of dice with your professor. We will call him Dr. Vikram. You pick a random number and compare it with Dr. Vikram's number which is `5`. If the difference between your number and Dr. Vikram's number is not larger than 1 then Dr. Vikram wins, if the difference is greater than 1 but not larger than 2, then you win! In all of other cases you noone wins. Your job is to write a program to simulate this game.

<!-- grade_id: cell-599bbbcdd28b7ae1, points: 1, grade: true, solution: true -->
```python
import random
r = random.randint(1,10)
print(r)

vikram_number = 5


if (vikram_number - r)<=1:
     print("vikram wins")
        
# YOUR CODE HERE
raise NotImplementedError()
```

## Lists and loops

By this stage, we have learned the basics of Python: variables, arithmetic operations, conditionals. However, sometimes we want to group data and explore a set of parameters, e.g. a set of pressure values in the tube. What 

```python
#lists

a_list = [1.0, 2.0, 3.0, 4.0]
print(a_list)
```

```python
# list numbering starts from 0 and goes until N-1
# to access its first element we use 
a_list[0]
# the third
a_list[2]
# the 4th
a_list[3]
# another way to get the last element is to use `-1`
a_list[-1]
# as you could have guessed, the second last element is [-2]
a_list[-2]

print("second last element = ", a_list[-2])
```

```python
# There are cool ways to access parts of an array:
# e.g. all the elements starting from the second one
a_list[1:]
```

### Loops

Now, to systematically access elements of a list we need to use loops. Python has two types of loops: `for` which goes through the elements of a list (actually they are generators in Python3, but lets not get too bookwarmy here) and `while` loop, but I won't mention it here for now.

```python
# Lets print all of the elemenet of the a_list 
for i in a_list:
    print(i)
```

```python
# What if I want to print the values of the a_list squared
for i in a_list:
    print(i**2)    
```

```python
# to get the length of the list we use the len function
len(a_list)
```

```python
# another way to work with loops is to go over its elements 
# for this we need to generate a list of indices
for i in range(len(a_list)):
    print(i)
```

### Project 6  

Do you remember Project 4, where we had to calculate the pressure of an ideal gas at certain temperature? Now I have a list of temperatures `t_list = [27, 60, 100, 150, 200]` and I want to calculate a list of pressures `P_list` corresponding to those temperatures. I remind that temperatures here are in `C`, there are two moles of gas, and volume of the vessel $V = 1 litre$".

## Functions

Ok ... hang in there! We are almost there!

Lets get familiar with Functions.

Why do we need them?

Well ... sometimes we want to separate blocks of code into a block of its own. For that we use functions.


```python
# Here is how we define Python functions

def name_of_a_func(input_param):
    print("myparam = ", input_param)
    return input_param

# here is how to call it

b = name_of_a_func(2)
print(b)
```

```python
# lets do a little project
# lets write a function that converts inches to cm (1 inch = 2.54 cm)

def inch_2_cm(param_inch):
    return param_inch*2.54

inch_2_cm(12)
```

```python
# Now lets write a function that convert temperature from degrees centigrade to Farenheit
#Tf = 1.8*Tc + 32 

def C_to_F(Tc):
    return Tc*1.8 + 32.

# lets call the function
C_to_F(0)
```

### Project 8

Now lets do something that will require almost all of the elements that we have learned today.
Lets calculate ... Fibonacci sequence! I remind you a Fibbonacci sequence 

$F_{n}=F_{n-1}+F_{n-2}$
with seed values: 
$F_{1}=0, F_{2}=1$

Lets write a function that calculates n-th element of the Fibbonacci sequence.

<!-- grade_id: cell-4719ec84e590e75b, points: 0, grade: false, solution: true -->
```python
nterm = 10
def Fib_sum(nterm):
#     f_sum 
#     f_sum = 0 + 1+ 1 + 2 + 3 ...
# if nterm = 0 => f_sum = 0
# if nterm = 1 => f_sum = 1
# nterm>1 => f_sum = sum (i =0, 1 ....[(n-1)+(n-2)])

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8418e848732a296b, points: 1, grade: true, solution: false -->
```python
assert_equal(Fib_sum(10),87)
```

## Bibliography:


http://www.science.uwaterloo.ca/~cchieh/cact/c120/idealgas.html

https://en.wikipedia.org/wiki/Fibonacci_number